```python
import numpy as np
import functools
from functools import reduce
from scipy.stats import kurtosis, skew
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
from scipy import stats
```

# Statistics with Functional Programming

In this exercise we will use imperative programming with Python to compute from a list of  N integers:

   - The mean - the first moment
   
   $\mu = \frac{1}{N} \sum_{i=1}^N x_i.$
        
        
  
   - The standard deviation - second moment
  
   $\sigma = \sqrt{\frac{1}{N} \sum_{i=1}^N (x_i - \mu)^2}$
        
        
        
   - The skew  - the third moment
   
   $\gamma_1 = \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^3$
        
        
        
   - The kurtosis - the fourth moment
  
   $\gamma_2 = \Big\{   \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^4  \Big\} - 3$
    

```python
def test_assert():
    """Test that assert is working.

    Args:
        None
    Returns:
        x with value 5
    """
    x = 5
    assert x == 5
    return x
    
y = test_assert() 
assert y == 5
```

# Define test data set X

```python
np.random.seed(5)
X = np.random.normal(size=100)
sns.set(color_codes=True)
sns.distplot(X)
print(X)

```

# First Moment

### Define a function, first_moment_term, to  compute the MEAN

 $\mu = \frac{1}{N} \sum_{i=1}^N x_i.$
 
 We will:
 
 1.  Map each term of the list to the correct moment
 2.  Sum the transformed elements over N
 

####  Define the transformation for the first moment terms 
    hint - it is the identity 

<!-- grade_id: cell-28bfdfc236ed5341, points: 0, grade: false, solution: true -->
```python
def first_moment_term(x):
    """Define the first moment term

    Args:
        x - An element of a list - a term
    Returns:
        transformed term - an identity in this case, x
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-87d4a0cf5b53adca, points: 10, grade: true, solution: false -->
```python
x = 1

new_term = first_moment_term(x)

print(new_term)

assert new_term == 1
```

<!-- grade_id: cell-8bc1d1d51a9571f8, points: 0, grade: false, solution: true -->
```python
def my_mean(my_list):
    """Define the first moment term

    Args:
        my_list - a list of integers
        
    Returns:
        the mean of the list
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-e29348aaca8d7264, points: 10, grade: true, solution: false -->
```python
#unit test
test_list = [1,2,3] 
mu = my_mean(test_list)
print(mu)
a = np.array([1,2,3])
assert mu == np.mean(a)
```

# Second Moment

### Define a function, second_moment_term, to compute the standard deviation

$\sigma = \sqrt{\frac{1}{N} \sum_{i=1}^N (x_i - \mu)^2}$
 
 We will:
 
 1.  Map each term of the list to the correct moment
 2.  Sum the transformed elements over N
 
 


####  Define the transformation for the first moment terms 

$(x_i - \mu)^2$
   

<!-- grade_id: cell-36c2bccf0cdc5e1e, points: 0, grade: false, solution: true -->
```python
def second_moment_term(x, mu):
    """Define the first moment term

    Args:
        x - An element of a list - a term
        mu - the mean
    Returns:
        transformed term - an identity in this case, (x - mu)**2
    """
    # YOUR CODE HERE
    raise NotImplementedError()
  
```

<!-- grade_id: cell-643b48e193aa2742, points: 10, grade: true, solution: false -->
```python
val = 5
mu = 3
term = second_moment_term(x, mu)
assert term == 4

```

### Transform and reduce to get the standard deviation


   $\sigma = \sqrt{\frac{1}{N} \sum_{i=1}^N (x_i - \mu)^2}$
   
    
   N = len(list)
   
   $\mu$ is the mean
   
   
##### Note:  Use Numpy for the square root  - 

      np.sqrt(x)
 

<!-- grade_id: cell-dbc8fbf2597f32bf, points: 0, grade: false, solution: true -->
```python
def my_stdev(my_list, mu):
    """Define the first moment term

    Args:
        my_list - a list of integers
        mu - the mean of the list
        
    Returns:
        the standard deviation of the list
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-f1f8efeff09c7446, points: 10, grade: true, solution: false -->
```python
test_list = [1,2,3] 
mu = my_mean(test_list)
print(mu)
sigma = my_stdev(test_list, mu)
print(sigma)
a = np.array([1,2,3])
assert mu == np.mean(a)
assert sigma == np.std(a)
print(np.std(a))
```

# Third Moment



## Define a function, third_moment_term, to compute the Skew

$\gamma_1 = \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^3$

We will:

Map each term of the list to the correct moment
Sum the transformed elements over N


<!-- grade_id: cell-e64bb8c186f0a559, points: 0, grade: false, solution: true -->
```python
def third_moment_term(x, mu, sigma):
    """Define the third moment term

    Args:
        x - An element of a list - a term
        mu - the mean
        sigma - the standard deviation
    Returns:
        transformed term - an identity in this case, ((x - mu)/sigma)**3
    """
    # YOUR CODE HERE
    raise NotImplementedError()
  
```

<!-- grade_id: cell-dca148de3b9ba00b, points: 10, grade: true, solution: false -->
```python
x = 5
mu = 3
sigma = 1
gamma_1 = third_moment_term(x, mu, sigma)
assert gamma_1 == 8
```

### Transform and reduce to obtain the skew


   $\gamma_1 = \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^3$
   
    
   N = len(list)
   
   $\mu$ is the mean
   
   
   $\sigma$ is the standard deviation
   
   

 

<!-- grade_id: cell-7464ccb7a0b68bed, points: 0, grade: false, solution: true -->
```python
def my_skew(my_list, mu, sigma):
    """Compute the third moment, the skew

    Args:
        my_list - a list of integers
        mu - the mean of the list
        sigma - the standard deviation
        
    Returns:
        the skew of the list
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-e236269fcf6cac89, points: 10, grade: true, solution: false -->
```python
test_list = [1,2,3] 
mu = my_mean(test_list)
print(mu)
sigma = my_stdev(test_list, mu)
print(sigma)
gamma_1 = my_skew(test_list, mu, sigma)
print(gamma_1)
a = np.array([1,2,3])
assert mu == np.mean(a)
assert sigma == np.std(a)
assert gamma_1 == skew(a)
print(skew(a))
```

# Fourth Moment

## Define a function, fourth_moment_term, to compute the kurtosis

$\gamma_2 = \Big\{   \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^4  \Big\} - 3$

We will:

Map each term of the list to the correct moment
Sum the transformed elements over N

<!-- grade_id: cell-edd9d9b8ea0b9666, points: 0, grade: false, solution: true -->
```python
def fourth_moment_term(x, mu, sigma):
    """Define the fourth moment term

    Args:
        x - An element of a list - a term
        mu - the mean
        sigma - the standard deviation
    Returns:
        transformed term - an identity in this case, ((x - mu)/sigma)**4
    """
    # YOUR CODE HERE
    raise NotImplementedError()
  
```

<!-- grade_id: cell-afb36679d5bf9073, points: 10, grade: true, solution: false -->
```python
x = 5
mu = 3
sigma = 1
gamma = fourth_moment_term(x, mu, sigma)
assert gamma == 16
```

 ### Transform and reduce to obtain the kurtosis

   $\gamma_2 = \Big\{   \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^4  \Big\} - 3$

    
   N = len(list)
   
   $\mu$ is the mean
   
   
   $\sigma$ is the standard deviation
   

<!-- grade_id: cell-e9dc7d41dab9a86f, points: 0, grade: false, solution: true -->
```python
def my_kurtosis(my_list, mu, sigma):
    """Define the fourth moment, the kurtosis

    Args:
        my_list - a list of integers
        mu - the mean of the list
        sigma - the standard deviation
        
    Returns:
        the skew of the list
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-ba6f2d63f0104426, points: 10, grade: true, solution: false -->
```python
test_list = [1,2,3] 
mu = my_mean(test_list)
print(mu)
sigma = my_stdev(test_list, mu)
print(sigma)
gamma_1 = my_skew(test_list, mu, sigma)
print(gamma_1)
gamma_2 = my_kurtosis(test_list, mu, sigma)
print(gamma_2)
a = np.array([1,2,3])
assert mu == np.mean(a)
assert sigma == np.std(a)
assert gamma_1 == skew(a)

assert gamma_2 - kurtosis(a) < np.finfo(np.float32).eps
#assert test_kurtosis == kurtosis(a)

```

### Sanity Check - compare with NUMPY

```python
from scipy.stats import kurtosis, skew
a = np.array([0,1,2,3,88,89,90,91])
print(np.mean(a))
print(np.std(a))
print(skew(a))
print(kurtosis(a))

```

```python
list1 = [0,1,2,3,88,89,90,91]
#sns.distplot(list1)
mu = my_mean(list1)
print(mu)
sigma = my_stdev(list1, mu)
print(sigma)
gamma_1 = my_skew(list1, mu, sigma)
gamma_2 = my_kurtosis(list1, mu, sigma)

print(mu)
print(sigma)
print(gamma_1)
print(gamma_2)
```

```python
print(np.finfo(np.float32).eps)
print(np.finfo(np.float).eps)
```

<!-- grade_id: cell-3479bc0c3a860bc2, points: 10, grade: true, solution: false -->
```python
from scipy.stats import kurtosis 
from scipy.stats import skew 
a = np.array([0,1,2,3,88,89,90,91])
assert mu - np.mean(a) < np.finfo(np.float).eps
assert sigma - np.std(a) < np.finfo(np.float).eps
assert gamma_1 - skew(a) < np.finfo(np.float).eps
assert gamma_2 - kurtosis(a) < np.finfo(np.float32).eps
#assert my_kurtosis1 - kurtosis(a) < np.finfo(np.float).eps
```

# Compute Stats for X 

```python
mu = my_mean(X)
print(mu)
sigma = my_stdev(X, mu)
print(sigma)
gamma_1 = my_skew(X, mu, sigma)
print(gamma_1)
gamma_2 = my_kurtosis(X, mu, sigma)
print(gamma_2)

```

<!-- grade_id: cell-31e2d4a5521239e7, points: 10, grade: true, solution: false -->
```python
a = np.array(list(X))
assert mu - np.mean(a) <= np.finfo(np.float).eps
assert sigma - np.std(a) <= np.finfo(np.float).eps
assert gamma_1 - skew(a) <= np.finfo(np.float).eps
assert gamma_2 - kurtosis(a) <= np.finfo(np.float32).eps
print("Congratulations you passed")
```

```python

```