```python
import numpy as np
import functools
from functools import reduce
from scipy.stats import kurtosis, skew
import pandas as pd
import seaborn as sns
import matplotlib.pyplot as plt
from scipy import stats
```

# Statistics with Imperative Programming

In this exercise we will use imperative programming with Python to compute from a list of  N integers:

   - The mean - the first moment
   
   $\mu = \frac{1}{N} \sum_{i=1}^N x_i.$
        
        
  
   - The standard deviation - second moment
  
   $\sigma = \sqrt{\frac{1}{N} \sum_{i=1}^N (x_i - \mu)^2}$
        
        
        
   - The skew  - the third moment
   
   $\gamma_1 = \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^3$
        
        
        
   - The kurtosis - the fourth moment
  
   $\gamma_2 = \Big\{   \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^4  \Big\} - 3$
        

```python
def test_assert():
    """Test that assert is working.

    Args:
        None
    Returns:
        x with value 5
    """
    x = 5
    assert x == 5
    return x
    
y = test_assert() 
assert y == 5
```

# Define test data set X

```python
# Python variables can be capital letters
np.random.seed(5)
X = np.random.normal(size=100)
sns.set(color_codes=True)
sns.distplot(X)
print(X)

```

# First Moment - the MEAN

### Define a function first_moment that computes the MEAN

 $\mu = \frac{1}{N} \sum_{i=1}^N x_i.$

<!-- grade_id: cell-d635eb21895c8eb3, points: 0, grade: false, solution: true -->
```python
def first_moment(list):
    """Compute the first moment - the mean.

    Args:
        list of integers of size n
    Returns:
        mean of the list
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-51fba0cc53017b71, points: 10, grade: true, solution: false -->
```python
l = [1,2]
mu = first_moment(l)
print(mu)

assert mu == 1.5
```

Make sure the empty list runs - N=0

<!-- grade_id: cell-17b978abc6e44f16, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-6634b61a65662b64, points: 5, grade: true, solution: false -->
```python
assert mu_empty == 0.0
```

# Second Moment - Standard Deviation

### Define a function second_moment that computes the Standard Deviation


   $\sigma = \sqrt{\frac{1}{N} \sum_{i=1}^N (x_i - \mu)^2}$
   
    
   N = len(list)
   
   $\mu$ is the mean
   
   
##### Note:  Use Numpy for the square root  - 

      np.sqrt(x)
 

<!-- grade_id: cell-7ad765bfd27b5d20, points: 0, grade: false, solution: true -->
```python
import numpy as np
def second_moment(list, mu):
    """Compute the second moment - the variance.

    Args:
        list of integers of size n
        mean - mean of the list
    Returns:
        second moment about the mean - the standard deviation
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
```

<!-- grade_id: cell-4490901cabe1c650, points: 10, grade: true, solution: false -->
```python
l = [1 , 2]
mu = first_moment(l)
sigma = second_moment(l, mu)
assert sigma == 0.5
print(sigma)
```

# Third Moment - the Skew

### Define a function, third_moment, that computes the Skew


   
   $\gamma_1 = \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^3$
   
   
  
   
   N = len(list)
   
   $\mu$ is the mean
   
   $\sigma$ is the standard deviation
   
   
 

<!-- grade_id: cell-1b7ec68c9705b5fa, points: 0, grade: false, solution: true -->
```python
import numpy as np
def third_moment(list, mu, sigma):
    """Compute the third moment - the skew.

    Args:
        list of integers of size n
        mu - mean  of the list
        sigma - standard deviation of the list
    Returns:
        third moment about the mean - the skew
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-1b01c15becd0529f, points: 10, grade: true, solution: false -->
```python
    
l = [1 , 2]
mu = first_moment(l)
sigma = second_moment(l, mu)
gamma_1 = third_moment(l, mu, sigma)
print(gamma_1)

assert gamma_1 == 0.0
```

# Fourth Moment

### Define a function, fourth_moment,  to compute the Kurtosis


   
   $\gamma_2 = \Big\{   \frac{1}{N} \sum_{i=1}^N \left[\frac{x_i - \mu}{\sigma}\right]^4  \Big\} - 3$
   
   
  
   
   N = len(list)
   
   $\mu$ is the mean
   
   $\sigma$ is the standard deviation
   

<!-- grade_id: cell-4dbec72aaba0c671, points: 0, grade: false, solution: true -->
```python
import numpy as np
def fourth_moment(list, mu, sigma):
    """Compute the fourth moment - the kurtosis.

    Args:
        list of integers of size n
        mu - mean of the list
        sigma - standard deviation of the list
    Returns:
        fourth moment about the mean - the skew
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-6e112b6609db7d58, points: 10, grade: true, solution: false -->
```python
l = [1 , 2]
gamma_2 = third_moment(l, mu, sigma)
print(gamma_2)

assert gamma_2 == 0
```

### Sanity Check - compare with NUMPY

```python
from scipy.stats import kurtosis, skew
a = np.array([0,1,2,3,88,89,90,91])
print(np.mean(a))
print(np.std(a))
print(skew(a))
print(kurtosis(a))

```

```python
list1 = [0,1,2,3,88,89,90,91]
mu = first_moment(list1)
sigma = second_moment(list1, mu)
gamma_1 = third_moment(list1, mu, sigma)
gamma_2 = fourth_moment(list1, mu, sigma)

print(mu)
print(sigma)
print(gamma_1)
print(gamma_2)
```

```python
print(np.finfo(np.float32).eps)
print(np.finfo(np.float).eps)
```

<!-- grade_id: cell-3012ff4e0398ad4a, points: 10, grade: true, solution: false -->
```python
from scipy.stats import kurtosis 
from scipy.stats import skew 
assert mu - np.mean(a) <= np.finfo(np.float).eps
assert sigma - np.std(a) <= np.finfo(np.float).eps
assert gamma_1 - skew(a) <= np.finfo(np.float32).eps
assert gamma_2 - kurtosis(a) <= np.finfo(np.float32).eps

```

# Compute Stats for X

```python
mu = first_moment(X)
print(mu)
sigma = second_moment(X, mu)
print(sigma)
gamma_1 = third_moment(X, mu, sigma)
print(gamma_1)
gamma_2 = fourth_moment(X, mu, sigma)
print(gamma_2)

```

<!-- grade_id: cell-7ee94dabad2c2f5f, points: 10, grade: true, solution: false -->
```python
a = np.array(list(X))
assert mu - np.mean(a) <= np.finfo(np.float).eps
assert sigma - np.std(a) <= np.finfo(np.float).eps
assert gamma_1 - skew(a) <= np.finfo(np.float).eps
assert gamma_2 - kurtosis(a) <= np.finfo(np.float32).eps
print("Congratulations you passed")
```

```python

```